Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (click the restart button in the tool bar or select Kernel$\rightarrow$Restart in the menu bar) and then **run all cells** (in the menu bar, select Cell$\rightarrow$Run All).

**Make sure you delete `raise NotImplementedError()`** (if existing) and fill in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Stephen Shell"
```

**Don't modify the formal TEST cells!**

---

# Core Language

For the first half of the homework, you are asked to write out some lines of code and define variables with particular names. In the second half of the homework, you are asked to fill code in functions.

*Don't forget:* You will need to delete the line with `raise NotImplementedError`, write your own solution, and then re-run the cell before running the test cell. Each time you change code in a cell, you  will need to re-run that cell before running any other cells that depend on it!

```python
# load the package
from nose.tools import assert_equal
```

### 1. (2 points) Create and use two variables

**1a.** (1 point) Create a variable called `var1` that equals the integer 20. Create another variable called `var2` that equals the float 3.5. Multiply `var1` by `var2` and call the result `var3`. 

<!-- grade_id: cell-6a3e3f6b30b2ee0e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
var1 = 20
var2 = 3.5
var3 = var1 * var2
print(type(var3))
# raise NotImplementedError()
```

Output:
```
<class 'float'>
```

<!-- grade_id: cell-6dff7ca5e48eb93c, points: 1, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
# Test the value of each variable
assert(var1 == 20)
assert(var2 == 3.5)
assert(var3 == 70)

print('Success!')
```

Output:
```
Success!
```

**1b.** (0.5 point) Write (using Markdown) what the type of `var3` should be and why.

<!-- grade_id: cell-341bf7ec6f32d4fb, points: 0.5, grade: true, solution: true -->
var3 should be a float, because, in python, whenever an integer is multiplied by a float, the answer becomes a float.

**1c.** (0.5 point) Check if `var3` is smaller than 15 and save result to the variable named `checkvar3`.

<!-- grade_id: cell-ae656e231c94fee7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
if var3 < 15:
    checkvar3 = True
else:
    checkvar3 = False
# raise NotImplementedError()
```

<!-- grade_id: cell-2f400f395a98fbd5, points: 0.5, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
assert(not checkvar3)

print('Success!')
```

Output:
```
Success!
```

### 2. (2 points) Operations on strings

**2a.** (1 point) Define two strings: `string1` contains the word "laser" and `string2` contains the word "printer". Combine the two words together with a space in between and call the resultant variable `printer`.

<!-- grade_id: cell-2ea540ee892203eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
string1 = "laser"
string2 = "printer"
printer = string1 + " " + string2
# raise NotImplementedError()
```

<!-- grade_id: cell-f06f27643c8964e6, points: 1, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
# test type of variables
assert(isinstance(string1, str))
assert(isinstance(string2, str))
assert(isinstance(printer, str))

# test final result
assert(printer == 'laser printer')

print('Success!')
```

Output:
```
Success!
```

**2b.** (1 point) Split the string by the letter "p" and save the resultant list in a variable named `reslist`.

<!-- grade_id: cell-4f3584f327c70d2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
reslist = printer.split('p')
# raise NotImplementedError()
```

<!-- grade_id: cell-8d90e109198c0c97, points: 1, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
assert(isinstance(reslist, list))
assert(reslist[0] == 'laser ')
assert(reslist[1] == 'rinter')

print('Success!')
```

Output:
```
Success!
```

### 3. (2 points) Make a list 

**3a.** (1 point) Create a variable called `alist` that contains, in order, the string 'apple', the boolean False, and the number 20. Write your answer in the following cell.

<!-- grade_id: cell-b8940cdbf0e33284, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
alist = ['apple', False, 20]
# raise NotImplementedError()
```

<!-- grade_id: cell-0b9aa531121ee3a4, points: 1, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
# This line tests to see if the variable you created, called `alist`, is of type list, and contains the correct element types
assert(isinstance(alist, list))
assert(isinstance(alist[0], str))
assert(isinstance(alist[1], bool))
assert(isinstance(alist[2], int))

# This line tests to make sure the correct elements are contained in the list
assert(alist[0] == 'apple')
assert(alist[1] == False)
assert(alist[2] == 20)

print('Success!')
```

Output:
```
Success!
```

**3b.** (0.5 point) Change the last entry of the list to equal the list `[3,5,9]`.

<!-- grade_id: cell-d21e7dc5ecf4f34f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
alist[-1] = [3,5,9]
# raise NotImplementedError()
```

<!-- grade_id: cell-efd6296b634211fd, points: 0.5, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
assert(isinstance(alist[2], list))
assert(alist[2] == [3,5,9])

print('Success!')
```

Output:
```
Success!
```

**3c.** (0.5 point) Check if the index 1 of `alist` is equivalent to the string "False" and assign the comparison result to variable `test`.

<!-- grade_id: cell-6126d685e7ef4b2c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
if alist[1] == "False":
    test = True
else:
    test = False
# raise NotImplementedError()
```

<!-- grade_id: cell-5956761dfba66c52, points: 0.5, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
assert(not test)

print('Success!')
```

Output:
```
Success!
```

### 4. (1 point) Make a dictionary 

Make a dictionary that contains the capital cities of a given state. Instead of all 50 states, just include Alabama, Texas, and California. The dictionary must be called `state_capitals` and should be able to used as follows (note: the capitial city should be capitalized).
    
    state_capitals['Texas'] = 'Austin'


<!-- grade_id: dictionary, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
state_capitals = {}
state_capitals.update({'Texas':'Austin', 'Alabama':'Montgomery', 'California':'Sacramento'})
# raise NotImplementedError()
```

<!-- grade_id: test_dictionary, points: 1, grade: true, solution: false -->
```python
""" Test code for the previous code. This cell should NOT give any errors when it is run."""
# makes sure your code returns a dictionary
assert(isinstance(state_capitals, dict))

# checks that the capitals can be accessed correctly
assert(state_capitals['Texas'] == 'Austin')
assert(state_capitals['Alabama'] == 'Montgomery')
assert(state_capitals['California'] == 'Sacramento')

# if nothing throws an assertion, good job!
print("Success!")
```

Output:
```
Success!
```

### 5. (2 points) Test for vowels 

Write a function that takes a character (i.e. a string of length 1) and returns True if it is a vowel, False otherwise. For example,

> isVowel("a")
  True <br>
> isVowel("b")
  False

You can compare the input character with a known list of vowels. Don't forget about case!

<!-- grade_id: isvowel, points: 0, grade: false, solution: true -->
```python
def isVowel(char):
    """Takes a character and returns True if it is a vowel, and False otherwise.
    
    Parameters
    ----------
    char : string
        string of length 1
        
    Returns
    -------
    boolean telling whether char is a vowel or not
    
    """
    # YOUR CODE HERE
    vowels = ['A', 'a', 'E', 'e', 'I', 'i', 'O', 'o', 'U', 'u']
    if len(char) > 1:
        return False
    else:
        if char in vowels:
            return True
        else:
            return False
    # raise NotImplementedError()
```

<!-- grade_id: test_isvowel, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""

assert isVowel('a')
assert not isVowel('m')
assert not isVowel(' ')

print("These tests pass, more in next cell.")
```

Output:
```
These tests pass, more in next cell.
```

<!-- grade_id: test_isvowel2, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""

assert isVowel('U')
assert not isVowel('This is not a vowel')

print("Success!")
```

Output:
```
Success!
```

### 6. (2 points) Reverse a string

Define a function reverse() that computes the reversal of a string. For example, 
    >>> reverse("I am testing")
    "gnitset ma I"

<!-- grade_id: reverse, points: 0, grade: false, solution: true -->
```python
def reverse(string):
    """Reverse string and then return the converted string.
    
    Parameters
    ----------
    string : str
        
        
    Returns
    -------
    string that is reversed
    """
    # YOUR CODE HERE
    reverse = string[::-1]
    return reverse
    # raise NotImplementedError()
```

<!-- grade_id: test_reverse, points: 2, grade: true, solution: false -->
```python
""" Test code for the previous function. This cell should NOT give any errors when it is run."""
assert_equal(reverse('I am testing'), 'gnitset ma I')
assert_equal(reverse('I am TesTinG'), 'GniTseT ma I')
assert_equal(reverse(''), '')

print("Success!")
```

Output:
```
Success!
```

# Graduate student problems below

Graduate students are **required** to complete problems 7 and 8 below. 

For undergraduates, you are encouraged to try. If you gain points in this graduate section, it can make up for the points you lose in the above undergraduate section. However, the maximum points you can get for this homework is the number of total points set up for undergraduates (in this homework, it is 9). For example, if you lose 2 points above and gain 2 points below, the final points you get is 9. If you lose 2 points above and gain 4 points below, the final points you get is still 9.

### 7. (2 points) Check for palindromes

Define a function `is_palindrome()` that recognizes palindromes (i.e. words and phrases that look the same written forwards and backwards). Make sure that case and spaces are ignored in your function; for example,
    >>> is_palindrome("Never odd or even") will return
    True
    
*Hint: You might need the .replace() method*

<!-- grade_id: palindrome, points: 0, grade: false, solution: true -->
```python
def is_palindrome(phrase):
    """Returns True for a palindrome and False otherwise.
    
    Parameters
    ----------
    phrase : string
        
        
    Returns
    -------
    Boolean
    """
    # YOUR CODE HERE
    phrase = phrase.lower()
    phrase = phrase.replace(' ', '')
    if phrase == phrase[::-1]:
        return True
    else:
        return False
    # raise NotImplementedError()
```

<!-- grade_id: test_palindrome, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""

assert is_palindrome("radar")
assert not is_palindrome("I am testing")

print("These tests pass, more in next cell.")
```

Output:
```
These tests pass, more in next cell.
```

<!-- grade_id: test_palindrome2, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""

assert is_palindrome("rad dar")
assert is_palindrome("Never odd or even")

print("Success!")
```

Output:
```
Success!
```

### 8. (2 points) Encode and decode using ROT-13

In cryptography, a Caesar cipher is a very simple encryption techniques in which each letter in the plain text is replaced by a letter some fixed number of positions down the alphabet. For example, with a shift of 3, A would be replaced by D, B would become E, and so on. The method is named after Julius Caesar, who used it to communicate with his generals. ROT-13 ("rotate by 13 places") is a widely used example of a Caesar cipher where the shift is 13. In Python, the encryption for ROT-13 may be represented by means of the following dictionary:

    rot_dict = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
           'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
           'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
           'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
           'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
           'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
           'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

Your task in this exercise is to implement an encoder/decoder of ROT-13. Once you're done, you will be able to read the following secret message:

       Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!

* Note that since English has 26 characters, your ROT-13 program will be able to both encode and decode texts written in English. You should be able to do this using the properties of a dictionary.


* Also note that your function should be able to deal with any special characters (not only those provided in the test cell, such as ? and !). **This is very important**


* The argument "letter" in the `rot13` function is only 1 single letter or 1 special character, not a string containing multiple letters or characters. So no need to use a for loop inside of the `rot13` function.


* You might need to use try-except or if-else.

<!-- grade_id: rot13, points: 0, grade: false, solution: true -->
```python
rot_dict = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
       'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
       'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
       'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
       'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
       'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
       'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

def rot13(letter):
    """Swap letter indicated in rot_dict.
    
    Parameters
    ----------
    letter : string of length 1
        letter to be swapped using the rot_dict dictionary
        
    Returns
    -------
    letter
    """
    # YOUR CODE HERE
    if letter in rot_dict.keys():
        return rot_dict.get(letter)
    else:
        return letter
    # raise NotImplementedError()
```

<!-- grade_id: test_rot13, points: 2, grade: true, solution: false -->
```python
""" Test code for the previous function. This cell should NOT give any errors when it is run."""
phrase = 'Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!'

translatedchars = list(map(rot13, phrase))  # map connects a function with 
                                            # each letter in phrase.
translated = ''.join(translatedchars)  # we got a bunch of characters back. Join combines them
                                       # into a string

# recode
retranslatedchars = list(map(rot13, translated))
retranslated = ''.join(retranslatedchars)
    
# Check decode
assert_equal(translated, "Caesar cipher? I much prefer Caesar salad!")
assert_equal(retranslated, phrase)  # check that it works both ways

print("Success!")
```

Output:
```
Success!
```

```python

```