Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (click the restart button in the tool bar or select Kernel$\rightarrow$Restart in the menu bar) and then **run all cells** (in the menu bar, select Cell$\rightarrow$Run All).

**Make sure you delete `raise NotImplementedError()`** (if existing) and fill in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Stephen Shell"
```

**Don't modify the formal TEST cells!**

---

# Core Language continued

```python
# import the package
from nose.tools import assert_equal
```

### 1. Compute the length of a given list or string (2 points)

Define a function that computes the length of a given list or string. (It is true that Python has the len() function built in, but do not use it for your solution — use a loop instead.)

For example:

    >>> length(['string1', 'string2'])
    2
    >>> length('string1')
    7

<!-- grade_id: cell-7cc34eb8c688c3a5, points: 0, grade: false, solution: true -->
```python
def length(inputs):
    """Compute the length of a given list or string, but without using `len`.
    
    Parameters
    ----------
    inputs : string or list
        
        
    Returns
    -------
    integer that is either the number of characters in the input string or the number of 
    entries in the input list.
    """
    # YOUR CODE HERE
    cnt = 0
    for i in inputs:
        cnt += 1
    return cnt
    # raise NotImplementedError()
```

<!-- grade_id: cell-3870df85ff34e2bb, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""

# Check a list
inputs = ['apple', 'orange']
assert_equal(length(inputs), 2)

# Check a string
inputs = 'tomato'
assert_equal(length(inputs), 6)

# What about a string with a space in it?
inputs = 'bacon and eggs'
assert_equal(length(inputs), 14)

# Also will look for not using the built-in len() function
len_old = len
len = None
inputs = 'tomato'
assert_equal(length(inputs), 6)
len = len_old

print("Success!")
```

Output:
```
Success!
```

### 2a-b. Convert list of words into list of integers of length of words (3 points)

Write a series of functions that maps a list of words into a list of integers representing the lengths of the corresponding words. Write them in two different ways: 

1. using a for-loop
1. using list comprehensions. 

Each function will return a list of word lengths, for example,

    >>> method1(['list', 'of', 'words'])
    [4, 2, 5]

<!-- grade_id: forloop, points: 0, grade: false, solution: true -->
```python
# part a
def method1(input1):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a for loop.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    # YOUR CODE HERE
    list1 = []
    for i in input1:
        list1.append(length(i))
    return list1
    # raise NotImplementedError()
```

<!-- grade_id: test_forloop, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method1(words), [3, 3, 5])
# Also will check that a for loop was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: listcomprehension, points: 0, grade: false, solution: true -->
```python
# part b
def method2(input2):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a list comprehension.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    # YOUR CODE HERE
    return [length(i) for i in input2]
    # raise NotImplementedError()
```

<!-- grade_id: test_listcomprehension, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method2(words), [3, 3, 5])
# Also will check that a list comprehension was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: test_all, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
words = ['five', 'eight', 'infinity']
# Check that the result is the same for all methods
assert_equal(method1(words), method2(words))
print('success!')
```

Output:
```
success!
```

### 3. Number of integers to reach 1 thousand (3 points)

Calculate how many integers, starting from 1 and counting each consecutive integer, are needed to add up to be more than 1 thousand. Use a `while` loop to do this.

<!-- grade_id: cell-3e995c9f072905fb, points: 0, grade: false, solution: true -->
```python
def ints_to_thousand():
    '''Find the number of consecutive integers (starting from 1) that are needed to add 
    together to sum to more than 1 thousand.'''
    
    # YOUR CODE HERE
    count = 0
    total = 0
    while not (total > 1000):
        count += 1
        total += count
    return count
    # raise NotImplementedError()
```

<!-- grade_id: cell-c17f04140621a1f3, points: 3, grade: true, solution: false -->
```python
"""(3 points) Test code for the previous function. This cell should NOT give any errors when it is run."""

# save the number of integers
numints = ints_to_thousand()

assert(numints == 45)

print('success!')
```

Output:
```
success!
```

### 4. Number lines in a text file (4 points)

Write a function, `num_lines()`, that numbers all of the lines in a text file, skipping empty lines, and writes the output to another specified file. For example, if the original text file contains:

    this is the first line: hello, python!
    this is a test line for second line
    
    this is a third line following the empty line
    

Then the output file should contain the original text but with the line number in the beginning. Note that for the empty line in the original file, it will still be an empty line in the output file (without line number!)

    0 this is the first line: hello, python!
    1 this is a test line for second line
    
    2 this is a third line following the empty line

The function uses inputfilename as the input file name argument and uses outputfilename as the output file name argument. **Hence, you don't need to create another filename inside of the function. Just use the default inputfilename and outpufilename**. In the test cell, we are using a file called `inputfile.txt` to serve as the input file. You can see its content by opening it using your laptop text editor. Note that this `inputfile.txt` file is sitting in the same directory as this notebook for you to read it easily using Python.

Tips: I highly recommend you make a copy of the original `inputfile.txt` file before you complete this problem. In previous semesters, some students will accidentally write new contents to the `inputfile.txt` file and the whole function will not work as expected. In case you think your answer is correct but the test cell still throws out errors, it might be due to the change of the original `inputfile.txt` file.

<!-- grade_id: lines, points: 0, grade: false, solution: true -->
```python
def num_lines(inputfilename, outputfilename):
    """Numbers all lines in a text file, skipping empty lines. Writes results to output file.
    
    Parameters
    ----------
    inputfilename : string
        Name of the input file to be read in.
        
    outputfilename : string
        Name of the output file to be written to.  
        
    Returns
    -------
    Nothing, but does write results to file outputfilename.
    """
    # YOUR CODE HERE
    count = 0
    in_file = open(inputfilename, "r")
    lines = in_file.readlines()
    with open(outputfilename, "w") as out_file:
        for in_line in lines:
            line_num = str(count)
            # check if line is empty
            if in_line in ['\n','\r\n']:
                # if empty, create an empty line
                out_line = "\n"
            else:
                # if not empty, add the line number and increase the count
                out_line = line_num + " " + in_line
                count += 1
            out_file.write(out_line)
        out_file.close()
    # raise NotImplementedError()
```

<!-- grade_id: test_lines, points: 4, grade: true, solution: false -->
```python
"""(4 points) Test code for the previous function. This cell should NOT give any errors when it is run."""

# run code
inputfilename = 'inputfile.txt'  # this is in the hw directory

# put output file in user directory:
import os
outputfilename = inputfilename.split('/')[-1].split('.')[0] + '_mod.txt'  # splitting up file name and removing path
outputfilename = '/'.join((os.getcwd(), outputfilename))  # this gives absolute path to user directory
num_lines(inputfilename, outputfilename)  # save to local directory

# check for existence of output file
assert os.path.exists(outputfilename)

# check for text in output file being correct
fin = open(outputfilename, 'r')  # prepare to read
lines = fin.readlines()
line = lines[0].split(' ')
assert_equal(line[0], '0')  # check first word in line
assert_equal(line[-1], 'eros.\n')  # check final word in line
assert_equal(lines[1], '\n')  # check empty line
line = lines[2].split(' ')
assert_equal(line[0], '1')  # check first word in line
assert_equal(line[-1], 'sodales.\n')  # check final word in line
assert_equal(lines[3], '\n')  # check empty line
line = lines[4].split(' ')
assert_equal(line[0], '2')  # check first word in line
assert_equal(line[-1], 'a.\n')  # check final word in line
assert_equal(lines[5], '\n')  # check empty line
line = lines[6].split(' ')
assert_equal(line[0], '3')  # check first word in line
assert_equal(line[-1], 'ut.\n')  # check final word in line
assert_equal(lines[7], '\n')  # check empty line
line = lines[8].split(' ')
assert_equal(line[0], '4')  # check first word in line
assert_equal(line[-1], 'malesuada.\n')  # check final word in line
fin.close()

print('Success!')
```

Output:
```
Success!
```

```python

```