Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (click the restart button in the tool bar or select Kernel$\rightarrow$Restart in the menu bar) and then **run all cells** (in the menu bar, select Cell$\rightarrow$Run All).

**Make sure you delete `raise NotImplementedError()`** (if existing) and fill in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Stephen Shell"
```

**Don't modify the formal TEST cells!**

---

These problems are to do some basic time indexing with pandas. The first code loads in some data from the USGS waterdata site.

```python
import numpy as np
import pandas as pd

url = 'https://waterdata.usgs.gov/tx/nwis/dv?format=rdb&site_no=08108700&referred_module=sw&begin_date=1993-01-01&end_date=2018-10-23'

# load the data
brazos = pd.read_csv(url, delim_whitespace=True, comment='#', skiprows=[34,], parse_dates=[2], index_col=2)
brazos.loc[:, '135330_00060_00003']  # mean daily discharge
```

Output:
```
datetime
1993-07-14        NaN
1993-07-15        NaN
1993-07-16        NaN
1993-07-17        NaN
1993-07-18        NaN
               ...   
2018-10-19    36700.0
2018-10-20    34900.0
2018-10-21    40300.0
2018-10-22    42100.0
2018-10-23    38300.0
Name: 135330_00060_00003, Length: 9233, dtype: float64
```

[3 points] Explain what each of the options in the function call to `pd.read_csv` does in the cell below. 

* Note that you need to  write down specifically what the parameter and the number are doing. For example, parse_dates: Tell pandas that the column specified should be imported as a date. In our case, that column index is 2, and is the 3rd column.


* Need to explicitly distinguish between **column index** vs. **position of column** (like in the above example, column index is 2 and it is equivalent to the 3rd column).

<!-- grade_id: cell-5bd73baa53a916d5, points: 3, grade: true, solution: true -->
**filepath_or_buffer       :** Gives function access to file. In this case, through a URL stored in the variable 'url'.

**delim_whitespace = True  :** Sets whitespace as a delimiter or not. In this case, true.

**comment = '#'            :** Skips specified lines. In this case, lines following a '#'.

**skiprows = $[34,]$       :** Skips line numbers (0-indexed) at the start of the file. In this case, lines 0-34 are skipped.

**parse_dates = $[2]$      :** Imports specified column as a date. In this case, column index 2.

**index_col = 2            :** Sets columns to be row labels. In this case, column index 2.

[2 points] Calculate the mean discharge in 2016 and 2017 and store the values in variables called 'brazos_mean_2016' and 'brazos_mean_2017'. Note that the discharge column name is "135330_00060_00003".

<!-- grade_id: cell-d5cbe7f06097fe2e, points: 0, grade: false, solution: true -->
```python
br = brazos.loc[:, '135330_00060_00003']

brazos_mean_2016 = br.loc['2016-01-01':'2016-12-31'].mean()
brazos_mean_2017 = br.loc['2017-01-01':'2017-12-31'].mean()

print(brazos_mean_2016)
print(brazos_mean_2017)

# raise NotImplementedError()
```

Output:
```
13075.825136612022
3443.3698630136987
```

<!-- grade_id: cell-356aa2cf22ec31e5, points: 2, grade: true, solution: false -->
```python
assert np.isclose(brazos_mean_2016, 13075.825136612022)
assert np.isclose(brazos_mean_2017, 3443.3698630136987)
```

[3 points] Compare early summer discharge (May, June, and July) with late summer/fall discharge (August, September, October) in 2016 and 2017. Store your variables with names:

   - brazos_mjj_2016
   - brazos_aso_2016
   - brazos_mjj_2017
   - brazos_aso_2017


<!-- grade_id: cell-838ea2db5fc2d9df, points: 0, grade: false, solution: true -->
```python
brazos_mjj_2016 = br.loc['2016-05-01':'2016-07-31'].mean()
brazos_aso_2016 = br.loc['2016-08-01':'2016-10-31'].mean()
brazos_mjj_2017 = br.loc['2017-05-01':'2017-07-31'].mean()
brazos_aso_2017 = br.loc['2017-08-01':'2017-10-31'].mean()

print(brazos_mjj_2016)
print(brazos_aso_2016)
print(brazos_mjj_2017)
print(brazos_aso_2017)

# raise NotImplementedError()
```

Output:
```
27803.695652173912
4584.945652173913
2802.8804347826085
1657.5217391304348
```

<!-- grade_id: cell-80a8857cc161af40, points: 3, grade: true, solution: false -->
```python
assert np.isclose(brazos_mjj_2016, 27803.695652173912)
assert np.isclose(brazos_aso_2016, 4584.945652173913)
assert np.isclose(brazos_mjj_2017, 2802.8804347826085)
assert np.isclose(brazos_aso_2017, 1657.5217391304348)
```

```python

```